```python
# Initialize Otter
import otter
grader = otter.Notebook("hw12.ipynb")
```

<img src="data8logo.png" alt="Data 8 Logo" style="width: 15%; float: right; padding: 1%; margin-right: 2%;"/>

# Homework 12: MLR

**Helpful Resource:**

- [Python Reference](http://data8.org/sp25/reference/): Cheat sheet of helpful array & table methods used in Data 8!

**Recommended Reading**: 

* [Classification](https://www.inferentialthinking.com/chapters/17/Classification.html)
* [Multiple Regression](https://inferentialthinking.com/chapters/17/6/Multiple_Regression.html)

**Credits**:

* **Conan Smallwood** + Fall 2024 Pedagogy Team
* **Dagny Streit** + Spring 2025 Pedagogy Team
* **Ella DeGuzman & Brandon Su** + Spring 2025 Content Team

Please complete this notebook by filling in the cells provided. **Before you begin, execute the cell below to setup the notebook by importing some helpful libraries.** Each time you start your server, you will need to execute this cell again.

For all problems that you must write explanations and sentences for, you **must** provide your answer in the designated space. Moreover, throughout this homework and all future ones, **please be sure to not re-assign variables throughout the notebook!** For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you thought you were passing previously!

**Deadline:**

This assignment is **due Sunday, 5/4 at 5:00pm PT**. Submissions after this time will be accepted for 24 hours and will incur a 20% penalty. Any submissions later than this 24 hour period will not be accepted unless an extension has been granted as per the [policies](http://data8.org/sp25/policies/) page. Turn it in by Saturday, 5/3 at 5:00pm PT for 5 extra credit points.

<span style="color:red">**Note: This homework has hidden tests on it. That means even though tests may say 100% passed, it doesn't mean your final grade will be 100%. We will be running more tests for correctness once everyone turns in the homework.**</span>

Directly sharing answers is not okay, but discussing problems with the course staff or with other students is encouraged. Refer to the [policies](http://data8.org/sp25/policies/#learning-cooperatively) page to learn more about how to learn cooperatively.

You should start early so that you have time to get help if you're stuck. Office hours are held Monday through Friday in [Warren Hall](https://www.berkeley.edu/map?warren) 101B. The office hours schedule appears [here](http://data8.org/sp25/schedule/).

---

The point breakdown for this assignment is given in the table below:
| Category | Points |
| --- | --- |
| Autograder (Coding questions) | 87 |
| Written | 13 |
| **Total** | 100 |

```python
# Run this cell to set up the notebook, but please don't change it.

import numpy as np
from datascience import * 
import hashlib

# These lines do some fancy plotting magic.
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
import warnings
warnings.simplefilter('ignore', FutureWarning)
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 1. Exploration of CO2 Emissions Dataset

In the following exploration, we will be working with a dataset containing 7385 cars, their specifications, and their CO2 emissions (g/km). 

Run the following cell to load information about `co2_raw` and preview the first few rows.

```python
co2_raw = Table().read_table("CO2 Emissions.csv").drop(2,5,6,9,10)
co2_raw.show(5)
```

For the sake of our exploration, let's rename some of the columns and call this new table `co2`.

```python
co2 = co2_raw.relabeled(
    "Engine Size(L)", "Engine").relabeled(
    "Fuel Consumption City (L/100 km)", "Fuel (City)").relabeled(
    "Fuel Consumption Hwy (L/100 km)", "Fuel (Hwy)").relabeled(
    "CO2 Emissions(g/km)", "Emissions")
co2.show(5)
```

We aim to predict the average CO2 emissions in grams per kilometer of an unseen car based on its __Engine__, __Cylinders__, __Fuel (City)__, and __Fuel (Hwy)__. To do this, we must first standardize and explore the data!

---

__Question 1.1__ Define a function `standardize` that takes in a column name `col` in a table `tbl`, and returns that column as an array in standard units. __(3 Points)__

```python
def standardize(tbl, col):
    arr = ...
    ...
```

```python
grader.check("q1_1")
```

---

__Question 1.2__ Define a function `correlation` that takes in the names of two columns (`x` and `y`) in a table `tbl`, and returns the correlation coefficient `r` between the two columns. __(3 Points)__

```python
def correlation(tbl, x, y):
    ...

correlation(co2, "Engine", "Emissions")
```

```python
grader.check("q1_2")
```

---

__Question 1.3__ Using the function you previously defined, calculate the correlation coefficient between `Emissions` and all other _numerical_ columns in the `co2` table. __(3 Points)__

Note: For the sake of passing the public tests, make sure `numerical_column_names` is an array of **strings**.

```python
numerical_column_names = ...

for col_name ...:
    current_r = ...
    print(f"{col_name}: {current_r}")
```

```python
grader.check("q1_3")
```

---

__Question 1.4__ Which numerical column is most strongly correlated with __Emissions__? Assign `highest_correlation` to the integer corresponding to the correct column. __(3 Points)__

1. Engine
2. Cylinders
3. Fuel (City)
4. Fuel (Hwy)

```python
highest_correlation = ...
```

```python
grader.check("q1_4")
```

<!-- BEGIN QUESTION -->

---

__Question 1.5__ Now, let's visualize the relationship between __Fuel (City)__ and __Emissions__. In the cell below, produce a scatter plot that plots the __Fuel (City)__ and __Emissions__ of every entry in the `co2` table with __Emissions__ on the y-axis. __(3 Points)__

```python
...
```

<!-- END QUESTION -->

Woah! It looks like the points are very tightly clustered around one line... or, two lines? What could be the cause of this? Let's take another look at our table and see which column could possibly explain what we see in the scatter plot.

---

__Question 1.6__ Assign `unique_cylinders_values` to a two-column table: the first column should contain unique value in `Cylinders`; the second column should contain the number of times that each of these unique values appear in the `co2` table. __(3 Points)__

```python
unique_cylinders_values = ...
unique_cylinders_values
```

```python
grader.check("q1_6")
```

It seems that most of the cars in the table fall under one of few __Cylinders__ values. Could this explain the split we see in the scatter plot? 

---

__Question 1.7__ Create a table called `co2_new` by adding a new column `Cylinders_and_City` to the `co2` table that is the sum of the columns `Fuel (City)` and `Cylinders`. Then, find the correlation between this new column and `Emissions`. __(3 Points)__

```python
co2_new = ...
strong_correlation = ...
strong_correlation
```

```python
grader.check("q1_7")
```

---

__Question 1.8__ True or False: The correlation coefficient between __Emissions__ and the sum of the __Cylinders__ and __Fuel (City)__ columns is _greater_ than the correlation coefficient between __Emissions__ and just the __Fuel (City)__ column. Assign `answer` to either `True` or `False`. __(3 Points)__

```python
answer = ...
```

```python
grader.check("q1_8")
```

Now, let's visualize this new column on a scatter plot.

```python
co2_new.scatter("Cylinders_and_City", "Emissions")
```

This looks more like a scatter plot that we're used to seeing! A lot of the time, there is more than one variable that contributes to the variable we are trying to predict. Sure, we can just use fuel consumption in the city to predict a car's emission levels, but there are other variables that we aren't accounting for! Including _multiple_ variables in our regression increases our model's ability to capture more complex patterns, and account for multiple factors that may be affecting the variable we want to predict!

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 2. Multiple Linear Regression to Predict CO2 Emissions

Before beginning this section of the homework, please read [this section of the textbook](https://inferentialthinking.com/chapters/17/6/Multiple_Regression.html) on Multiple Linear Regression!

What we've done so far in this class is called Simple Linear Regression (SLR), which involves one attribute to estimate a single numerical outcome. Precisely, the equation for the regression line has the following form: $$y = \text{slope} \cdot x+\text{intercept}$$ Here, $x$ is the attribute, and we predict the value of $y$ by applying a slope and adding an intercept.

Now, we will extend this idea to Multiple Linear Regression (MLR), which uses multiple attributes instead of just one. We make the prediction by multiplying each attribute value by its corresponding slope and summing them all together. The equation for the regression line has the following form: $$y = \text{slope}_1 \cdot x_1 + \text{slope}_2 \cdot x_2+ ... +\text{slope}_k \cdot x_k$$ for $k$ predictor variables. Here, $x_1$, $x_2$, $\cdots$, $x_k$ represent different numerical predictor variables from the dataset, and we still predict a single numerical output $y$. Each predictor variable gets its own slope. Note that in this homework, our equation for the regression line will not have an intercept, but in general, the equation for the regression line can have an intercept!

We use MLR when we believe more than one attribute influences the outcome we are trying to predict. To build the model, we split the dataset into a training set and a testing set. We use the training set to find the best slopes, and then evaluate its performance using Root Mean Squared Error (RMSE) on the testing set.

In this part of the homework, we will use MLR to predict the CO2 emissions of a car using its engine size, cylinders, city fuel consumption, and highway fuel consumption.

---

__Question 2.1__ In order to accurately assess how well a model can generalize and predict unseen data, it is important to first divide your dataset into a training set and a test set. Fill in the following code to first select the columns for engine size, cylinders, city fuel consumption, highway fuel consumption, and emissions. Then, to shuffle the data, reserve the first $1000$ data points for the test set, and assign the rest to the training set. __(6 Points)__

Note: Do __not__ use `tbl.split` (you will not pass the autograder). Take a look at [Section 17.2.2](https://inferentialthinking.com/chapters/17/2/Training_and_Testing.html#generating-a-test-set) on how to get your training and test sets.

Hint: `numerical_co2` should be a table of only the numerical columns from `co2`.

```python
np.random.seed(8) #DO NOT CHANGE

numerical_co2 = ...
total_rows = ...
shuffled_tbl = ...

test_set = ...
training_set = ...
```

```python
grader.check("q2_1")
```

Now the question is: How do we obtain the slopes ($\text{slope}_1, \text{slope}_2, ... \text{slope}_k$) that minimizes the RMSE (reminder that the RMSE helps quantify the average size of the error our prediction model makes on any given testing observation). Well, our end result is going to look something like `minimize(fit)` where `fit` is a function that takes in an array of slopes and returns the RMSE of our predictions. However, there is a lot that goes into this calculation; let's take this one step at a time.

---

__Question 2.2__ Define the function `predict` that takes in an array of slopes and a row object of attributes, and returns a number representing the corresponding prediction. You may assume that the array of slopes and the row object of attributes will be of the same length. __(3 Points)__

Hint 1: Take a look at [Section 17.6.2](https://inferentialthinking.com/chapters/17/6/Multiple_Regression.html#multiple-linear-regression)

Hint 2: You can perform array arithmetic on row objects as well!

```python
def predict(slopes, attributes):
    ...
```

```python
grader.check("q2_2")
```

---

__Question 2.3__ Now that we've defined a function that returns a prediction given an array of slopes and a row object of attributes, let's define `rmse` which takes in an array of slopes, a _table_ of attributes that we would like to use in our prediction, and an array of the _true_ values of y we want to predict. The function should return the RMSE of our predictions. __(10 Points)__

Hint: Take a look at [Section 17.3](https://inferentialthinking.com/chapters/17/3/Rows_of_Tables.html) for a review on row objects

```python
def rmse(slopes, attributes, y):
    ...
    for ...:
        ...
    return ...
```

Let's test out the `rmse` function we just defined on our training set with some made up slope values that definitely are not optimal.

```python
train_attributes = training_set.drop("Emissions")
train_emissions = training_set.column("Emissions")

fake_slopes = make_array(1,1,1,1)
rmse(fake_slopes, train_attributes, train_emissions)
```

```python
grader.check("q2_3")
```

Those are some pretty big errors! It's quite clear that the optimal array of slopes is not just a bunch of 1's. However, now that we've defined `rmse`, we are finally ready to define `train_fit`!

---

__Question 2.4__ Define `train_fit`, which takes in an array of slopes and returns the RMSE on the <u>training set</u>. __(5 Points)__

Hint: We're looking to predict emissions from our attributes!

```python
def train_fit(slopes_array):
    ...
```

```python
grader.check("q2_4")
```

---

__Question 2.5__ Obtain the array of slopes that minimize the RMSE on the training set, and assign that to `best_slopes`. __(4 Points)__

*Hint*: Make sure to use `fake_slopes` as one of your arguments. When the numerical optimization begins, the model might try out some inoptimal slopes first as it's searching for what will eventually be the right ones -- meaning, the ones that minimize RMSE!

Note: Take a look at [Section 17.6.2](https://inferentialthinking.com/chapters/17/6/Multiple_Regression.html#multiple-linear-regression)

```python
#This cell might take around one minute to run
best_slopes = ...
```

```python
grader.check("q2_5")
```

Let's see how well our optimal model performs on the training set.

```python
train_fit(best_slopes)
```

Not bad at all! 

Let's take a look at the values in `best_slopes`

```python
best_slopes
```

That seems a bit weird... doesn't it? Let's look at `train_attributes` again, the table of attributes we use to predict an emissions value.

```python
train_attributes.show(5)
```

You might have noticed something surprising: even though `Engine` and `Emissions` had a strong positive correlation (with an r value of around 0.85), the slope value for `Engine` obtained in MLR, which is denoted by the first element in the `best_slopes` array, is negative.

This happens when the `Engine` attribute is correlated with other attributes in the model!

---

__Question 2.6__ Now that we've obtained the array of slope values that minimize the RMSE, it's time to assess the performance of our model on the test set! Assign `test_rmse` to the RMSE obtained by using our regression line on the test set. __(3 Points)__

```python
test_attributes = ...
test_emissions = ...

test_rmse = ...
test_rmse
```

```python
grader.check("q2_6")
```

<!-- BEGIN QUESTION -->

---

__Question 2.7__ Compare the values of the training RMSE and testing RMSE. What do the values suggest about our model's ability to perform well on unseen data? __(5 Points)__ 

*Hint*: If the testing RMSE was much larger than the training RMSE, what would this say about our model's ability to perform well on unseen data?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

Including many attributes in your Multiple Linear Regression can increase your model's ability to capture complex patterns, but at the same time, including _too_ many features can make your model prone to overfitting. 

---

Suppose that we're restricted to using _only three attributes_ for our predictions. Let's try _removing_ the `Fuel (City)` attribute, and only using `Engine`, `Cylinders`, and `Fuel (Hwy)` in our predictions!

```python
def train_rmse_nocity(slopes):
    return rmse(slopes, train_attributes.drop("Fuel (City)"), train_emissions)

best_slopes_nocity = minimize(train_rmse_nocity, start = make_array(1,1,1), smooth = True, array = True)
best_slopes_nocity
```

```python
train_nocity_rmse = rmse(best_slopes_nocity, train_attributes.drop("Fuel (City)"), train_emissions)
test_nocity_rmse = rmse(best_slopes_nocity, test_attributes.drop("Fuel (City)"), test_emissions)
print(f"Training RMSE excluding Fuel (City): {train_nocity_rmse}")
print(f"Test RMSE excluding Fuel (City): {test_nocity_rmse}")
```

<!-- BEGIN QUESTION -->

---

__Question 2.8__ Notice how both the Training RMSE and Test RMSE barely increased, even though we fully omitted one attribute from our predictions! Why could have removing `Fuel (City)` from our predictions led to a negligible decrease in the accuracy of our predictions? __(5 Points)__

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 3. Exploration of SwingVision Dataset

In this next exploration, we will be working with data from SwingVision that contains information about tennis shots hit by players. [SwingVision](https://swing.vision/) is a startup founded by previous Data 8 professor Swupnil Sahai that analyzes tennis and pickleball players' forms and playing statistics, providing feedback to improve their game!

Run the following cell to load information about the dataset and preview the first few rows.

```python
Table().read_table("shots_cleaned.csv").show(5)
```

For this exploration, we'll look at a cleaned version of this dataset and assign it to `shots`. Note that "hit" refers to when the ball is hit by the racket, and "bounce" refers to when the ball bounces on the court.

```python
shots = Table().read_table("shots_cleaned.csv").select("match_id", 
                                                   "set_id", 
                                                   "game_id", 
                                                   "hit_type", 
                                                   "hit_court_side", 
                                                   "hit_velocity",
                                                   "bounce_court_side", 
                                                   "bounce_location_long", 
                                                   "bounce_location")
shots.show(5)
```

---

__Question 3.1__ In a game of tennis, keeping constant pressure on the opponent is very important. There's a lot of data in the table, let's limit our analysis to the following shots:

* Shots that were considered as "ground_stroke"
* Shots that were hit by the player closer to the camera (i.e., "near")
* Shots that bounced on the other side of the court (i.e., "far")
* Shots that were _not_ out.

Name this new table `ground_strokes`, and only keep the columns `hit_velocity`, `bounce_location_long`, `bounce_location`. __(5 Points)__

```python
np.random.seed(8) #DO NOT CHANGE

ground_strokes = ...

ground_strokes = ground_strokes.sample(with_replacement = False).take(np.arange(1000))
ground_strokes.show(5)
```

```python
grader.check("q3_1")
```

---

__Question 3.2__ Assign `unique_bounce_locations` to a two-column table: the first column should contain each unique value in `bounce_location_long`; the second column should contain the number of times that each of these unique values appear in the `ground_strokes` table. __(3 Points)__

```python
unqiue_bounce_locations = ...
unqiue_bounce_locations
```

```python
grader.check("q3_2")
```

Notice that the _hit_velocity_ and _bounce_location_ columns don't contain ordinary integers or strings like we're used to seeing. In fact, each value in _hit_velocity_ and _bounce_location_ is a _list_ of coordinates (x, y, z)! If you don't know what a list is, you can think of it as something similar to an array. Let's add columns to the `ground_strokes` table that contains the _individual coordinates_ as floats.

```python
#Run this cell
import ast
def extract_x(lst):
    return ast.literal_eval(lst)[0]

def extract_y(lst):
    return ast.literal_eval(lst)[1]

def extract_z(lst):
    return ast.literal_eval(lst)[2]

hit_velocity_x = make_array()
hit_velocity_y = make_array()
hit_velocity_z = make_array()

for coordinates in ground_strokes.column("hit_velocity"):
    hit_velocity_x = np.append(hit_velocity_x, extract_x(coordinates))
    hit_velocity_y = np.append(hit_velocity_y, extract_y(coordinates))
    hit_velocity_z = np.append(hit_velocity_z, extract_z(coordinates))

bounce_x = make_array()
bounce_y = make_array()
bounce_z = make_array()

for coordinates in ground_strokes.column("bounce_location"):
    bounce_x = np.append(bounce_x, extract_x(coordinates))
    bounce_y = np.append(bounce_y, extract_y(coordinates))
    bounce_z = np.append(bounce_z, extract_z(coordinates))

coords = ground_strokes.with_columns(
    "hit_velocity_x", hit_velocity_x, "hit_velocity_y", hit_velocity_y, "hit_velocity_z", hit_velocity_z).with_columns(
    "bounce_x", bounce_x, "bounce_y", bounce_y, "bounce_z", bounce_z).drop("hit_velocity", "bounce_location").where(
    "hit_velocity_z", are.below(15))

coords
```

In this exploration, we want to predict or classify how _deep_ a shot lands, based on the initial velocity of the ball when it is hit by a player. To describe the coordinates (x,y,z):

* _x_ describes left/right (right being positive)
* _y_ describes forward/backward (forward being positive)
* _z_ describes up/down (up being positive)

```python
#Run this cell
colors = {'service_box': 'orange', 'no_mans_land': 'blue'}
color_list = [colors[group] for group in coords['bounce_location_long']]
plt.scatter(data = coords.where('bounce_location_long', 'service_box'), x = "hit_velocity_y", y = "hit_velocity_z", s = 2, color = 'orange', label='service_box')
plt.scatter(data = coords.where('bounce_location_long', 'no_mans_land'), x = "hit_velocity_y", y = "hit_velocity_z", s = 2, color = 'blue', label='no_mans_land')
plt.xlabel("Hit Velocity (y coordinate)")
plt.ylabel("Hit Velocity (z coordinate)")
plt.legend();
```

We can see a fairly clear decision boundary between shots that land in "no_mans_land", colored in blue, and shots that land in the "service_box", colored in orange.

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 4. k-Nearest Regression with SwingVision Dataset

For a refresher on kNN, review [Chapter 17](https://inferentialthinking.com/chapters/17/Classification.html) from the textbook!

kNN regression is a method used to predict a numerical value based on nearby data points. To make a prediction, we first calculate the Euclidean distance between the point we are trying to predict and each point in the data set. We then find the k closest points, known as the nearest neighbors. Instead of taking a majority vote like we did in kNN classificaiton, we take the average of the values we are trying to predict from the nearest neighbors. This average value becomes our prediction.

Sure, we can classify (with some degree of accuracy) whether a shot we hit will land in the _service_box_ or _no_mans_land_ based on it's initial velocity, but maybe we would want to be more specific about how _deep_ our shot will land on the opponent's side. After all, the deeper the shot lands, the more pressure you put on the opponent! This is where kNN regression is useful — rather than labeling the bounce location, we can use it to predict the depth (numerical position) where the ball is likely to bounce.

---

__Question 4.1__ Again, let's split our `coords` dataset into a training set and a test set, reserving the first $200$ shots for the test set and putting the rest into the training set. __(6 Points)__

Note: Do __not__ use `tbl.split`.

```python
np.random.seed(8) #DO NOT CHANGE

num_shots = ...
shuffled_coords = ...

test_coords = ...
training_coords = ...
```

```python
grader.check("q4_1")
```

---

__Question 4.2__ As always, let's define a function that calculates the distance between two points according to the relevant attributes. Complete the `distance` function that takes in two rows, and calculates the Euclidean distance using `hit_velocity_y` and `hit_velocity_z`. __(3 Points)__

Note: The formula for Euclidean distance is $\sqrt{(x_1 - x_2)^2 + (y_1 - y_2)^2}$

```python
def distance(row1, row2):
    y_diff = ...
    z_diff = ...
    ...
```

```python
grader.check("q4_2")
```

---

__Question 4.3__ Define a function `predict_nn` that takes in a row, and returns the predicted _y_ coordinate of where the ball landed based on the k nearest neighbors. __(10 Points)__

*Hint*: Make sure to read the introduction of Question 4 carefully before approaching this problem!

```python
def predict_nn(point, k):
    num_training = ...
    dists = ...
    for ...:
        ...
    train_dists = ...
    return ...
```

```python
grader.check("q4_3")
```

---
__Question 4.4__ Now, perform kNN regression with $k = 7$ on the test set defined earlier. Calculate the RMSE of our predictions. __(8 Points)__

```python
nn_predictions = ...
num_test = ...
...

knn_regression_rmse = ...
knn_regression_rmse
```

```python
grader.check("q4_4")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 5. **(OPTIONAL)** kNN Classification with SwingVision Dataset

Now, we will develop a k-Nearest Neighbors (kNN) classifier to classify how deep a shot lands by predicting whether it bounces in the "service_box" or in "no_mans_land", using the ball's inital velocity in the forward/backward (`hit_velocity_y`) and up/down directions (`hit_velocity_z`).

---

__Question 5.1__ Define the function `classify` that takes in a point (a row object) and `k`, and classifies that point using the data in the training set.

*Hint*: Take a look at [Section 17.4.5](https://inferentialthinking.com/chapters/17/4/Implementing_the_Classifier.html#implementation-step-1) on how to implement this function.

Note: You don't need to follow the skeleton code to answer this question!

```python
def classify(point, k):
    num_training = ...
    dists = ...
    for ...:
        ...
    train_dists = ...
    return ...
```

```python
grader.check("q5_1")
```

Let's test our classifier on the first data point in our test set with $k = 7$!

```python
classify(test_coords.row(0), 7)
```

```python
test_coords.show(1)
```

---

__Question 5.2__ Define a function `accuracy` which takes in an array of predictions (an array of actual values) and returns the proportion of correct predictions.

```python
def accuracy(preds, y):
    ...
```

```python
grader.check("q5_2")
```

---

__Question 5.3__ Now, determine the accuracy of the kNN classifier we developed on the test set we defined earlier, using $k = 7$.

**Hint**: Consider using a for loop to store the prediction of each shot from `test_coords` in the array `predictions`.

```python
num_test = ...
predictions = ...
...

accuracy(predictions, test_coords.column("bounce_location_long"))
```

```python
grader.check("q5_3")
```

You're done with Homework 12!

## Pets of Data 8

Boba thinks you should treat yourself to a boba!

<img src="boba.jpeg" width="40%" alt="Shorthair cat smirking at camera"/>

Congrats on finishing the final homework!

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Written Work Submission

Below, you will see two cells. Running the first cell will automatically generate a PDF of all questions that need to be manually graded, and running the second cell will automatically generate a zip with your autograded answers. You are responsible for submitting both the coding portion (the zip) and the written portion (the PDF) to their respective Gradescope portals. **Please save before exporting!**

> **Important: You must correctly assign the pages of your PDF after you submit to the correct gradescope assignment. If your pages are not correctly assigned and/or not in the correct PDF format by the deadline, we reserve the right to award no points for your written work.**

If there are issues with automatically generating the PDF in the first cell, you can try downloading the notebook as a PDF by clicking on `File -> Save and Export Notebook As... -> Webpdf`. If that doesn't work either, you can manually take screenshots of your answers to the manually graded questions and submit one single PDF of your screenshots. Either way, **you are responsible for ensuring your submission follows our requirements, we will NOT be granting regrade requests for submissions that don't follow instructions.**

**You must submit the PDF generated via one of these methods, we will not accept screenshots or Word documents.**

```python
from otter.export import export_notebook
from os import path
from IPython.display import display, HTML
name = 'hw12'
export_notebook(f"{name}.ipynb", filtering=True, pagebreaks=True)
if(path.exists(f'{name}.pdf')):
    display(HTML(f"Download your PDF <a href='{name}.pdf' download>here</a>."))
else:
    print("\n Pdf generation failed, please try the other methods described above")
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1_1 (points: [0, 0, 3])
```
>>> isinstance(standardize(co2, 'Engine'), np.ndarray)
True
```
```
>>> type(standardize(co2, 'Engine').item(0)) in set([float, np.float64])
True
```
```
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> get_hash(np.round(standardize(co2, 'Engine').item(0), 4)) == '6eb876585034f98e94657ca747d7d1cc'
True
```


Tests for q1_2 (points: [0, 3])
```
>>> -1 <= correlation(co2, 'Engine', 'Emissions') <= 1
True
```
```
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> get_hash(np.round(correlation(co2, 'Engine', 'Emissions'), 4)) == '3a02549881dc0cb05bc1d7baf1685fff'
True
```


Tests for q1_3 (points: [0, 0, 3])
```
>>> -1 <= correlation(co2, 'Engine', 'Emissions') <= 1
True
```
```
>>> np.all(numerical_column_names == make_array('Engine', 'Cylinders', 'Fuel (City)', 'Fuel (Hwy)'))
True
```
```
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> get_hash(np.round(current_r, 4)) == '99156cd6931a2455506ab6d5c31b62a9'
True
```


Tests for q1_4 (points: [0])
```
>>> type(highest_correlation) == int and 1 <= highest_correlation <= 4
True
```


Tests for q1_6 (points: [0, 0, 0, 3])
```
>>> unique_cylinders_values.num_rows == 8
True
```
```
>>> unique_cylinders_values.column(1).item(3)
2446
```
```
>>> unique_cylinders_values.column(0).item(0)
3
```
```
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> get_hash(unique_cylinders_values.column(1).item(4)) == '28fc2782ea7ef51c1104ccf7b9bea13d'
True
```


Tests for q1_7 (points: [0, 1.5, 1.5])
```
>>> 'Cylinders_and_City' in co2_new.labels
True
```
```
>>> co2_new.column('Cylinders_and_City').item(4)
18.1
```
```
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> get_hash(np.round(strong_correlation, 4)) == '79591d2930dca681c78d688dfc32af32'
True
```


Tests for q1_8 (points: [0])
```
>>> type(answer) == bool
True
```


Tests for q2_1 (points: [0, 0, 0, 0, 6])
```
>>> np.all(numerical_co2.labels == make_array('Engine', 'Cylinders', 'Fuel (City)', 'Fuel (Hwy)', 'Emissions'))
True
```
```
>>> np.all(test_set.labels == make_array('Engine', 'Cylinders', 'Fuel (City)', 'Fuel (Hwy)', 'Emissions'))
True
```
```
>>> np.all(test_set.labels == make_array('Engine', 'Cylinders', 'Fuel (City)', 'Fuel (Hwy)', 'Emissions'))
True
```
```
>>> shuffled_tbl.num_rows == 7385
True
```
```
>>> test_set.column('Fuel (Hwy)').item(0) == 10.1
True
```


Tests for q2_2 (points: [0, 3])
```
>>> slopes = make_array(1, 2, 3)
>>> attributes = make_array(4, 5, 6)
>>> predict(slopes, attributes) == 32
True
```
```
>>> slopes = make_array(0, 1, 2)
>>> attributes = make_array(7, 8, 9)
>>> predict(slopes, attributes) == 26
True
```


Tests for q2_3 (points: [0, 10])
```
>>> type(rmse(fake_slopes, train_attributes, train_emissions)) in set([float, np.float64])
True
```
```
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> get_hash(np.round(rmse(fake_slopes, train_attributes, train_emissions), 4)) == '476915be7ef34cd4334ccaed88c212b2'
True
```


Tests for q2_4 (points: [0, 5])
```
>>> type(train_fit(make_array(8, 8, 8, 8))) in set([float, np.float64])
True
```
```
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> get_hash(np.round(train_fit(make_array(8, 8, 8, 8)), 4)) == '2084462bf54dcd204bf26206428c01c4'
True
```


Tests for q2_5 (points: [1, 1, 2])
```
>>> np.isclose(train_fit(best_slopes), 22.753604146634835)
True
```
```
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> get_hash(np.round(best_slopes.item(0), 4)) == 'fcf6a9732c9e8eabb3f99b5d2de6c92e'
True
```
```
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> get_hash(np.round(best_slopes.item(1), 4)) == '69fca06fd6ab7330b32d310716051dba'
True
```


Tests for q2_6 (points: [0, 3])
```
>>> type(test_rmse) in set([float, np.float64])
True
```
```
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> get_hash(np.round(test_rmse, 4)) == 'd45b282417ae86b295a7bfa3edaad8b1'
True
```


Tests for q3_1 (points: [0, 2, 3])
```
>>> ground_strokes.num_rows == 1000
True
```
```
>>> np.all(ground_strokes.labels == make_array('hit_velocity', 'bounce_location_long', 'bounce_location'))
True
```
```
>>> ground_strokes.column(1).item(4)
'no_mans_land'
```


Tests for q3_2 (points: [0, 3])
```
>>> unqiue_bounce_locations.num_rows == 2
True
```
```
>>> unqiue_bounce_locations.column('count').item(0) + unqiue_bounce_locations.column('count').item(1) == 1000
True
```


Tests for q4_1 (points: [0, 6])
```
>>> shuffled_coords.num_rows == 996
True
```
```
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> get_hash(np.round(test_coords.column('hit_velocity_x').item(0), 4)) == '3ba81a508ad7c81fbe46780facf9d3a5'
True
```


Tests for q4_2 (points: [3])
```
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> get_hash(np.round(distance(training_coords.row(0), training_coords.row(1)), 4)) == '9728b30cb5dca6c44fd8293f804a5fbe'
True
```


Tests for q4_3 (points: [0, 10])
```
>>> type(predict_nn(test_coords.row(0), 3)) in set([float, np.float64])
True
```
```
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> get_hash(np.round(predict_nn(test_coords.row(0), 3), 4)) == 'cdcada915d0258f9b9e7b37252418a00'
True
```


Tests for q4_4 (points: [0, 8])
```
>>> len(nn_predictions) == 200
True
```
```
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> get_hash(np.round(knn_regression_rmse, 4)) == '26e36142112bb6ff28252841573b1c0f'
True
```


Tests for q5_1 (points: [0, 0])
```
>>> type(classify(test_coords.row(0), 7)) == str
True
```
```
>>> classify(test_coords.row(0), 7)
'no_mans_land'
```


Tests for q5_2 (points: [0, 0, 0])
```
>>> type(accuracy(make_array(4, 3, 2), make_array(1, 2, 4))) in set([float, np.float64])
True
```
```
>>> 0 <= accuracy(make_array(4, 3, 2), make_array(1, 3, 4)) <= 1
True
```
```
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> get_hash(np.round(accuracy(make_array(4, 3, 2), make_array(1, 3, 4)), 4)) == 'd80cd8a9bc59ca0c1f479e4eca3c64be'
True
```


Tests for q5_3 (points: [0, 0])
```
>>> 0 <= accuracy(predictions, test_coords.column('bounce_location_long')) <= 1
True
```
```
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> get_hash(np.round(accuracy(predictions, test_coords.column('bounce_location_long')), 4)) == '57eeec0a6974ecb4e9fcf68fab052f7b'
True
```
